# Lab 02: Iocane Powder

## Overview

> Man in Black: All right. Where is the poison? The battle of wits has begun. It ends when you decide and we both drink, and find out who is right... and who is dead. 

The line above is from the perennial favorite 1980s movie adaptation of William Goldman's *The Princess Bride*, wherein a mysterious hero sits down to a battle of wits with a villainous Sicilian kidnapper. The setup: two cups positioned between the two, one of which (purportedly) contains a colorless, odorless, lethal poison (viz., iocane powder). After a guess is made as to which cup contains the poison, both drink, and the winner is the one left standing.

For this machine problem you will write a program that simulates multiple rounds of this battle of wits, allowing the player to repeatedly guess which cup is poisoned. The computer will "place" the poison before the player guesses, and
will reveal who is right... and who is dead, afterwards.

At the outset, the computer will always place the poison in cup 2 before letting the player guess, but after enough guesses have been entered the computer will start to place the poison based on the pattern of previous guesses so as to outsmart the player.

Here's a sample game session (note how the silly player keeps alternating guesses, and that the computer catches on to this fact after a while):

    Where is the iocane powder: my cup (1) or yours (2)? 1
    Wrong! Ha! Never bet against a Sicilian!

    You died 1 times, and I drank the poison 0 times

    Where is the iocane powder: my cup (1) or yours (2)? 2
    Good guess! Ack! I drank the poison!

    You died 1 times, and I drank the poison 1 times

    Where is the iocane powder: my cup (1) or yours (2)? 1
    Wrong! Ha! Never bet against a Sicilian!

    You died 2 times, and I drank the poison 1 times

    Where is the iocane powder: my cup (1) or yours (2)? 2
    Good guess! Ack! I drank the poison!

    You died 2 times, and I drank the poison 2 times

    Where is the iocane powder: my cup (1) or yours (2)? 1
    Wrong! Ha! Never bet against a Sicilian!

    You died 3 times, and I drank the poison 2 times

    Where is the iocane powder: my cup (1) or yours (2)? 2
    Wrong! Ha! Never bet against a Sicilian!

    You died 4 times, and I drank the poison 2 times

    Where is the iocane powder: my cup (1) or yours (2)? 1
    Wrong! Ha! Never bet against a Sicilian!

    You died 5 times, and I drank the poison 2 times

    Where is the iocane powder: my cup (1) or yours (2)? 2
    Wrong! Ha! Never bet against a Sicilian!

    You died 6 times, and I drank the poison 2 times

    Where is the iocane powder: my cup (1) or yours (2)? 1
    Wrong! Ha! Never bet against a Sicilian!

    You died 7 times, and I drank the poison 2 times


## Implementation

To keep track of the pattern of previous guesses, you will use a dictionary that maps a pattern (of fixed length) to a list of counts for the subsequent guess. 

For instance, imagine that the computer observes the player continuing to alternate guesses across ten separate attempts, like so: '1', '2', '1', '2', '1', '2', '1', '2', '1', '2'. If we are using a pattern detection length of three, then after the fourth guess we can create an entry in our dictionary that maps the key '121' to the list [0, 1], where the second value (1) in the list indicates that the player guessed '2' following the sequence '1', '2', '1'. After the fifth guess, we create the entry '212' &rarr; [1, 0], and after the sixth guess we update the value for '121' to [0, 2] (since the user guesses '2' again, after the sequence '1', '2', '1').

Once the player enters a series of guesses that matches a previously seen pattern, the computer should place the poison in the cup that the player is *least likely to guess next*. When the player enters the next guess, the dictionary should be updated to reflect the actual guess.

This means that if the computer has yet to see a given pattern of guesses, or when the counts are tied, it will have to place the poison "blindly" --- your implementation should simply place the poison furthest away from itself (cup 2). 

### `record_guess`

The first function you are to complete is `record_guess`. It will take the following arguments:

- a dictionary to update (possibly containing previously recorded pattern &rarr; list mappings)
- a pattern string
- a guess -- which is either '1' or '2'.  

If necessary, the function will create a new entry for the pattern (if one doesn't already exist), then record the updated count for the guess. Since the dictionary is updated in place (i.e., mutated), the function will not return anything. 

Complete the function below, checking your implementation with the test cases that follow when you're ready. Note that in the future, the bulk of the description for functions we ask you to implement will simply be placed in the functions' docstrings, as below.

*Hints: the [`int`](https://docs.python.org/3/library/functions.html#int) function can be used to convert strings to integers, and you might find the dictionary's [`setdefault`](https://docs.python.org/3/library/stdtypes.html?highlight=setdefault#dict.setdefault) method useful.*

<!-- grade_id: record_guess, points: 0, grade: false, solution: true -->
```python
def record_guess(pattern_dict, pattern, guess):
    """Updates the `pattern_dict` dictionary by either creating a new entry
    or updating an existing entry for key `pattern`, increasing the count 
    corresponding to `guess` in the list."""
    if not pattern in pattern_dict:
        pattern_dict[pattern] = [0, 0]
        
    if guess == '1':
        pattern_dict[pattern][0] += 1
    elif guess == '2':
        pattern_dict[pattern][1] += 1
        
    
```

<!-- grade_id: test_record_guess, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()
d = {}
record_guess(d, '121', '1')
tc.assertDictEqual(d, {'121': [1, 0]})
record_guess(d, '222', '2')
record_guess(d, '121', '1')
tc.assertDictEqual(d, {'121': [2, 0], '222': [0, 1]})
record_guess(d, '122', '2')
record_guess(d, '121', '2')
record_guess(d, '222', '2')
tc.assertDictEqual(d, {'121': [2, 1], '122': [0, 1], '222': [0, 2]})
```

### `next_placement`

The next function you'll write will take a dictionary of pattern &rarr; counts mappings and a string representing the pattern of most recent guesses, and return the next best location (either '1' or '2') for the poison (i.e., to try and outwit the player). If the pattern hasn't been seen previously or the counts are tied, the function should return '2'.

<!-- grade_id: next_placement, points: 0, grade: false, solution: true -->
```python
def next_placement(pattern_dict, pattern):
    if not pattern in pattern_dict:
        return '2'
    elif int(pattern_dict[pattern][0]) >= int(pattern_dict[pattern][1]):
        return '2'
    else:
        return '1'
            
```

<!-- grade_id: test_next_placement, points: 5, grade: true, solution: false -->
```python
# (5 points)

from unittest import TestCase
tc = TestCase()
tc.assertEqual(next_placement({}, '121'), '2')
tc.assertEqual(next_placement({'121': [2, 0]}, '121'), '2')
tc.assertEqual(next_placement({'121': [2, 5]}, '121'), '1')
tc.assertEqual(next_placement({'121': [2, 5]}, '212'), '2')
tc.assertEqual(next_placement({'121': [5, 5]}, '121'), '2')
tc.assertEqual(next_placement({'121': [15, 5]}, '121'), '2')
tc.assertEqual(next_placement({'121': [2, 5],
                               '212': [1, 1]}, '212'), '2')
tc.assertEqual(next_placement({'121': [2, 5],
                               '212': [1, 3]}, '212'), '1')
```

### `play_interactive`

Now for the fun bit. The function `play_interactive` will take just one argument --- the length of patterns to use as keys in the dictionary --- and will start an interactive game session, reading either '1' or '2' from the player as guesses, using the functions you wrote above and producing output as shown in the sample game session at the beginning of this writeup. If the player types in any other input (besides '1' or '2'), the game should terminate.

*Hint: the [`input`](https://docs.python.org/3/library/functions.html#input) function can be used to read input from the user as a string.*

<!-- grade_id: play_interactive, points: 0, grade: true, solution: true -->
```python
def play_interactive(pattern_length=4):
    nums = [""]*pattern_length
    right = 0
    wrong = 0
    pats = {}
    comp_guess = "0"
    while True:
        guess = input("Where is the iocane powder: my cup (1) or yours (2)?")
        if guess != '1' and guess != '2':
            break
        if right + wrong < pattern_length:
            nums[right + wrong] = guess
        if right + wrong < pattern_length or comp_guess == "2":
            if guess == "1":
                wrong += 1
                print ("Wrong! Ha! Never bet against a Sicilian!")
            else:
                right += 1
                print("Good guess! Ack! I drank the poison!")
            print ("\n") 
            print ("You died " + str(wrong) + " times, and I drank the poison " + str(right) + " times")
            print ("\n")    
        if comp_guess == "1":
            if guess == "2":
                wrong += 1
                print ("Wrong! Ha! Never bet against a Sicilian!")
            else:
                right += 1
                print("Good guess! Ack! I drank the poison!")
            print ("\n") 
            print ("You died " + str(wrong) + " times, and I drank the poison " + str(right) + " times")
            print ("\n") 
        if right + wrong > pattern_length:
            temp = [""]*pattern_length
            for i in range(3):
                temp[i] = nums[i+1]
            temp[3] = guess
            nums = temp
        if right + wrong >= pattern_length:
            pattern = ""
            for i in nums:
                pattern += i
            record_guess(pats, pattern, guess)
            comp_guess = next_placement(pats, pattern)
            
            
play_interactive()        
```

Output:
```
Where is the iocane powder: my cup (1) or yours (2)?2
Good guess! Ack! I drank the poison!


You died 0 times, and I drank the poison 1 times


Where is the iocane powder: my cup (1) or yours (2)?1
Wrong! Ha! Never bet against a Sicilian!


You died 1 times, and I drank the poison 1 times


Where is the iocane powder: my cup (1) or yours (2)?2
Good guess! Ack! I drank the poison!


You died 1 times, and I drank the poison 2 times


Where is the iocane powder: my cup (1) or yours (2)?
```

## `play_batch`

Finally, so that we can check your implementation against a lengthier sequence of guesses without having to play an interactive session, implement the `play_batch` function, which will take the `pattern_length` argument as your `play_interactive` function did, but will also take a sequence of guesses. The function will return the total numbers of wins and losses, as determined by the same algorithm as before.

<!-- grade_id: play_batch, points: 0, grade: false, solution: true -->
```python
def play_batch(guesses, pattern_length=4):
    nums = [""]*pattern_length
    right = 0
    wrong = 0
    pats = {}
    tracker = 0
    comp_guess = "2"
    plays = 0
    for i in guesses:
        pattern = ""
        guess = i
        if comp_guess == "2":
            if guess == "1":
                wrong += 1
            elif guess == "2":
                right += 1
        if comp_guess == "1":
            if guess == "1":
                right += 1
            elif guess == "2":
                wrong += 1
        tracker += 1
        if tracker <= pattern_length:
            nums[tracker - 1] = guess
        if tracker > pattern_length:
            for i in range (pattern_length):
                pattern = pattern + nums[i]
            record_guess(pats, pattern, guess)
            temp = [""]*pattern_length
            for i in range (pattern_length - 1):
                temp[i] = nums[i + 1]
            temp[pattern_length - 1] = guess
            nums = temp
            pattern = ""
            for i in range (pattern_length):
                pattern = pattern + nums[i]
            comp_guess = next_placement(pats, pattern)
        #print (right, wrong)
    return(right, wrong) 
```

<!-- grade_id: test_play_batch, points: 10, grade: true, solution: false -->
```python
# (10 points)

from unittest import TestCase
tc = TestCase()
tc.assertEqual(play_batch(['1', '1', '1', '1', '1', '1'], 3), (0, 6))
tc.assertEqual(play_batch(['1', '2', '1', '2', '1', '2'], 3), (2, 4))
tc.assertEqual(play_batch(['1', '2', '1', '2', '1', '2'], 4), (3, 3))
tc.assertEqual(play_batch(['1', '2'] * 100, 5), (3, 197))
tc.assertEqual(play_batch(['1', '1', '2', '1', '2', '1'] * 100, 2), (398, 202))
tc.assertEqual(play_batch(['1', '1', '2', '1', '2', '1'] * 100, 3), (201, 399))
tc.assertEqual(play_batch(['1', '1', '2', '1', '2', '1'] * 100, 5), (4, 596))
import random
random.seed(0, version=2)
tc.assertEqual(play_batch((random.choice(['1', '2']) for _ in range(10000)), 4), (5047, 4953))
```

Output:
```
0 1
0 2
0 3
0 4
0 5
0 6
0 1
1 1
1 2
2 2
2 3
2 4
0 1
1 1
1 2
2 2
2 3
3 3
0 1
1 1
1 2
2 2
2 3
3 3
3 4
3 5
3 6
3 7
3 8
3 9
3 10
3 11
3 12
3 13
3 14
3 15
3 16
3 17
3 18
3 19
3 20
3 21
3 22
3 23
3 24
3 25
3 26
3 27
3 28
3 29
3 30
3 31
3 32
3 33
3 34
3 35
3 36
3 37
3 38
3 39
3 40
3 41
3 42
3 43
3 44
3 45
3 46
3 47
3 48
3 49
3 50
3 51
3 52
3 53
3 54
3 55
3 56
3 57
3 58
3 59
3 60
3 61
3 62
3 63
3 64
3 65
3 66
3 67
3 68
3 69
3 70
3 71
3 72
3 73
3 74
3 75
3 76
3 77
3 78
3 79
3 80
3 81
3 82
3 83
3 84
3 85
3 86
3 87
3 88
3 89
3 90
3 91
3 92
3 93
3 94
3 95
3 96
3 97
3 98
3 99
3 100
3 101
3 102
3 103
3 104
3 105
3 106
3 107
3 108
3 109
3 110
3 111
3 112
3 113
3 114
3 115
3 116
3 117
3 118
3 119
3 120
3 121
3 122
3 123
3 124
3 125
3 126
3 127
3 128
3 129
3 130
3 131
3 132
3 133
3 134
3 135
3 136
3 137
3 138
3 139
3 140
3 141
3 142
3 143
3 144
3 145
3 146
3 147
3 148
3 149
3 150
3 151
3 152
3 153
3 154
3 155
3 156
3 157
3 158
3 159
3 160
3 161
3 162
3 163
3 164
3 165
3 166
3 167
3 168
3 169
3 170
3 17
```

```python

```